# Homework 1  (15 pts)
### Objectives
1. This homework makes use of python variables, expression, lists and numpy arrays.  
1. Indexing into arrays is an important skill to acquire by the end of the homework.  
### In order for automatic grading and validation to work, it is very important that you make use of the **exact variable names** specified in the problem statement.  
### I will always place these **variable** names that must be used for automatic grading in **bold** 
### Let's get started! First thing to do is always import the modules you need. So run the next line of code. 

```python
import numpy as np 
```

1. Create a variable named **My_First_Variable** and set its value to 2.  (1 pt)

<!-- grade_id: cell-2745506731a213ef, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
My_First_Variable = 2
### END SOLUTION
```

### We test solutions using assert statements.  For fun, try to do this problem correctly and incorrectly and see what it does when you run the code block below, containing an `assert` statement.  

<!-- grade_id: cell-a1e9fdf4667bd8d6, points: 1, grade: true, solution: false -->
```python
assert My_First_Variable == 2
```

2. Create a variable **v** and set its value equal to 0.5. Compute the variable **w** using the following equation $\textbf{w} = 2\frac{1-\textbf{v}}{1+\textbf{v}}$ 
* A useful technique here is to make use of parenthesis to control the order of operations. If you place (1-v) and and (1+v) inside parenthesis, they will be operated on first, before dividing.  
* You can also feel free to make intermediate variables e.g., x = 1-v and y = 1+v if that helps you organize the solution.  There is no perfect version of the code.  

<!-- grade_id: cell-989cfd2c6b242c38, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
v = 0.5
w = 2*(1-v)/(1+v)
### END SOLUTION
```

<!-- grade_id: cell-2d6405744fe5ae36, points: 1, grade: true, solution: false -->
```python
assert v == 0.5
assert w == 2/3
```

3. Create a list named **days** containing the days of the week, starting on 'Monday' and ending on 'Sunday' 

<!-- grade_id: cell-f5f3ef2fa3517087, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
days = ['Monday','Tuesday','Wednesday','Thursday','Friday','Saturday', 'Sunday']
### END SOLUTION
```

<!-- grade_id: cell-e97c023c69ed93a6, points: 1, grade: true, solution: false -->
```python
assert len(days) == 7
assert days[0] == 'Monday'
assert days[6] == 'Sunday'
### BEGIN HIDDEN TESTS
assert days[1] == 'Tuesday'
assert days[2] == 'Wednesday'
assert days[3] == 'Thursday'
assert days[4] == 'Friday'
assert days[5] == 'Saturday'
### END HIDDEN TESTS
```

4. Use indexing to create two new lists, **Weekend** which contain 'Saturday' and 'Sunday', and **Weekdays** which contain the other 5 working days of the week. Do not solve this problem by manually making another list.  The goal here is to think about indexing. 

<!-- grade_id: cell-cd3015af9040c09b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
Weekend = days[5:7]
Weekdays = days[0:5]
### END SOLUTION
```

<!-- grade_id: cell-1794b59379065020, points: 1, grade: true, solution: false -->
```python
assert len(Weekend) == 2
```

<!-- grade_id: cell-507689ec563fd49b, points: 1, grade: true, solution: false -->
```python
assert len(Weekdays) == 5
```

5. Given the list **words** below, use indexing, string concatentation, to fill the list **greeting** by appending the phrases, ["Hello Sir", "Dear Friend", "Friend Price"].  Pay attention to the spaces in the phrases - a space is a string with a space inside it, ' '.   

<!-- grade_id: cell-aa293f5d1282fa6b, points: 0, grade: false, solution: true -->
```python
words = ['Hello', 'Friend', 'Sir', 'Dear','Price']
greeting = list()
### BEGIN SOLUTION 
space = ' '
greeting.append(words[0]+space+words[2])
greeting.append(words[3]+space+words[1])
greeting.append(words[1]+space+words[4])
### END SOLUTION
```

<!-- grade_id: cell-b1c2c2249824c112, points: 2, grade: true, solution: false -->
```python
assert len(greeting) == 3
```

6. Create a variable **v** which is a list that contains the values 7, 16, 15, 14, 22.  Create a variable **w** which is a list that contains the values -1,-0.5,0,0.5,1 

<!-- grade_id: cell-612596199a906073, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
v = [7,16,15,14,22]
w = [-1,-0.5,0,0.5,1]
### END SOLUTION
```

<!-- grade_id: cell-2161a504819ff4c6, points: 1, grade: true, solution: false -->
```python
assert len(v) == 5
assert len(w) == 5
assert v[2] == 15
assert w[2] == 0 
```

7. Create a variable z calculated by multiplying the 4th element of **v** with the 2nd element of **w**.  Use indexing into the arrays **v** and **w** to solve the problem.     

<!-- grade_id: cell-cc9a25d65da6209b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
z = v[3]*w[1]
### END SOLUTION
```

<!-- grade_id: cell-9a81ae0f43d19e8a, points: 1, grade: true, solution: false -->
```python
assert z == -7
```

8. Create a new list **y** which contains the first 3 elements of **v** and a new list **z** which containes the middle 3 entries (2nd to 4th) of **w**.  Do not just simply "figure it out in your head" and write out a new list definition.  Use indexing to copy the correct elements of **v** into **y** and **w** into **z**. 


<!-- grade_id: cell-2754de3f28bea4ab, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
y = v[0:3]
z = w[1:4]
### END SOLUTION
```

<!-- grade_id: cell-6d2a8307c4674c13, points: 1, grade: true, solution: false -->
```python
assert len(y) == 3
assert len(z) == 3
```

9. Merge **y** and **z** to make a new list **yz** which contains the entries of **y** followed by the entries of **z**


<!-- grade_id: cell-18de3a19fd8b0db6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
yz = y+z
### END SOLUTION
```

```python
assert yz[3] == w[1]
```

10. Using negative indexing, make a new variable **last_vw** which is a new list constructed by merging the last *2* elements of **v** followed by the last *4* elements of **w**

<!-- grade_id: cell-e677d07d4021a5c4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
last2_v = v[-2:]
last4_w = w[-4:]
last_vw = last2_v+last4_w
### END SOLUTION
print(last_vw)
```

Output:
```
[14, 22, -0.5, 0, 0.5, 1]
```

<!-- grade_id: cell-bdf650196df2ebf6, points: 1, grade: true, solution: false -->
```python
assert last_vw[0] == 14
assert last_vw[-1] == 1
### BEGIN HIDDEN TESTS
assert last_vw[3] == 0
### END HIDDEN TESTS
```

11. Convert the lists **v** and **w** into a numpy array *with the same variable name* using the `array` function.  Compute a new array **vxw** by multiplying **v** and **w**.  

<!-- grade_id: cell-d20078385bad627e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
v = np.array(v)
w = np.array(w)
vxw = v*w
### END SOLUTION
```

<!-- grade_id: cell-84861554104a9652, points: 1, grade: true, solution: false -->
```python
assert vxw[0] == -7
assert vxw[-1] == 22 
```

12. Create a numpy array named **z** with 10 integers, all with value -100 

<!-- grade_id: cell-97c7f75e071d15d6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
z = -100*np.ones(10)
### END SOLUTION
```

<!-- grade_id: cell-17dffd810b651887, points: 1, grade: true, solution: false -->
```python
assert z.size == 10
assert z[0] == -100
```

13. Make an array named **B** containing elements in the list [5,7,9,12,13,15,17,19]. Change the value of the 3rd to 5th element to -100. You should be able to do this in a single line of code by index slicing.    

<!-- grade_id: cell-53428fb44cc641d5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
B = np.array([5,7,9,12,13,15,17,19])
B[2:5] = -100
### END SOLUTION
```

<!-- grade_id: cell-6c2c0206f55515c0, points: 1, grade: true, solution: false -->
```python
assert B[2] == -100
assert B[4] == -100
```

14.  Create an array named **digit_2** containing the whole numbers from 10 to 99, including 99

<!-- grade_id: cell-61f67ad6eae1b1dd, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
digit_2 = np.arange(10,100,1)
### END SOLUTION
```

<!-- grade_id: cell-fb5af02e1e74dfdc, points: 1, grade: true, solution: false -->
```python
assert digit_2[0] == 10
assert digit_2[-1] == 99
```