Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Homework 3 

* Practice some of the fundamentals of writing and using for loops. 
* Make some basic line using `matplotlib` and the `pyplot` submodule



```python
import numpy as np
from matplotlib import pyplot as plt
from numpy import random
```

1. Write a function **cummean**, which computes the cumulative mean of any array provided as an input argument **numbers**.   The function returns an array whose 
* first element is the mean of the first element of **numbers**.  
* second element is the mean of the first 2 elements of **numbers**
* The third element is the mean of the first 3 elements of **numbers** 
 .... (keep going)
* last element is the mean of all the elements of **numbers** 

The function should return an array of cumulative means where each element adds one more entry of **numbers** to the mean.    

<!-- grade_id: cell-cd812e0e45fad54d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3c6235b56fcd687c, points: 5, grade: true, solution: false -->
```python
rng = random.default_rng(seed = 1112)
numbers = rng.uniform(0,1,1000)
means = cummean(numbers)
assert np.isclose(means[3],0.64,rtol = 0.01)
assert np.isclose(means[99],0.43,rtol = 0.01)
assert np.isclose(means[499],0.49,rtol = 0.01)
```

2. Grading Problem. You get to write my code to make the final grades in this class. 

Your grade is based on 3 categories - 
* Classwork 40% 
* Tests 45%
* Homework 15%

As an example here are the grades for 5 students in each category

* classwork = [80, 95, 85, 55, 95]
* tests = [75, 90,90, 70,95]
* homework = [100, 90,95,90,95]

The grading scale is 

* A: finalgrade >= 89   
* B: finalgrade >= 79   
* C: finalgrade >= 69   
* D: finalgrade >= 59
* F: final grade < 60    

Write a function that takes in lists (or arrays) **classwork**, **tests**, **homework**, and returns the final student letter grades.   

<!-- grade_id: cell-d2b2deefa1a04793, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-a5969987b9b779a9, points: 5, grade: true, solution: false -->
```python
classwork = np.array([80, 95, 85, 70, 95])
tests = np.array([75, 90,90, 55,95])
homework = np.array([100, 90,95,90,95])
lettergrade = makegrades(classwork,tests,homework)
print(lettergrade)
assert len(lettergrade) == 5
assert lettergrade[0] == 'B'
assert lettergrade[-2] == 'D'
```

 3. In your first year on a job you earn a starting salary.  *Every other year* your salary goes up by a fraction increase.  Make a function **future_salary** that returns an array that has your salary for every year up to a number of years provided as an argument to the function.  The input arguments to the function are the starting salary, increase (as a fraction) and number of years.   

<!-- grade_id: cell-23c4db18cee6cf3b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-afca405ddb47060d, points: 4, grade: true, solution: false -->
```python
nyears = 5
increase = 0.1 
starting = 100000
salary = future_salary(starting,increase,nyears)
assert len(salary)==nyears
assert salary[0] == salary[1]
assert salary[-1] > salary[-2]
assert np.isclose(salary[2],110000,rtol=0.0001)
```

<!-- grade_id: cell-005e37989b442134, points: 4, grade: true, solution: false -->
```python
nyears = 6
increase = 0.15 
starting = 40000
salary = future_salary(starting,increase,nyears)
assert len(salary)==nyears
assert salary[0] == salary[1]
assert salary[-1] == salary[-2]
assert np.isclose(salary[-1],52900,rtol=0.0001)
```

4. The sigmoid function is invaluable in both psychology and neuroscience research and also in artificial intelligence.
The equation for a sigmoid function is
$$ s(x) = \frac{e^{ax}}{e^{ax}+1} $$
a is a parameter that determines the shape of the function.

Write a function **sigmoid** that can be used to compute the sigmoid function following the equation above.

The argument **x** is an array containing the values at which you want to know the values of **s**.

The function should be defined such that **a** *takes the default value of 1*.

The function should take input arguments **x** and **a** and return **s** 

Useful hints: 
* The exponential function can be called as np.exp().  
  

<!-- grade_id: cell-66d0aa05f9d1903c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-7e345d7c79fc0e4b, points: 3, grade: true, solution: false -->
```python
x = np.arange(-10,10.5,0.5)
s = sigmoid(x)
assert len(s) == len(x)
assert s[20] == 0.5
assert np.isclose(s[19],0.3775,rtol = 0.005)
```

<!-- grade_id: cell-8ee4f4069074d441, points: 3, grade: true, solution: false -->
```python
x = np.arange(-5,5,0.5)
a = 0.5
s = sigmoid(x,a) 
assert len(s) == len(x) 
assert s[10] == 0.5
assert np.isclose(s[5],0.223,rtol = 0.005)
```

5. Visualize how the value of **a** changes the shape of the sigmoid. 
* Make a graph of the values of **s** for three different values of **a** - 0.5, 1.0, 1.5. 
* Let **x** range from -10 to 10 in steps of 0.5.  
* Use different colored lines for the 3 graphs, and make a legend labeling the lines by the value of a.    
* You may use your own artistic sensibilities to choose colors and symbols. 

Colors - 
* blue - 'b',
* green - 'g',
* red - 'r',
* cyan - 'c',
* magenta - 'm',
* yellow - 'y',
* black - 'k',
* white - 'w' 

Line styles 

* '-' solid line 
* '--' dashed line
* '-.' dash dot line
* ':'  dotted line 

Marker shapes 
* '.'- point
* 'o' - circle
* 'x' - x marker 
* 'D' = diamond marker
* 'H' - hexagon marker 
* 's' - square marker
* '+' plus marker 

The example shown here contains all the important features of a well made graph.

![alt text](sigmoid.png "Sigmoid plot")


<!-- grade_id: cell-f97228f75effc113, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE

```